```python
# Initialize Otter
import otter
grader = otter.Notebook("26100319_2.ipynb")
```

# **Assignment 1 (Part 2)**

<h3>Assignment Instructions</h3>


* **How to install 'otter'**: Run `pip install otter-grader` in your Anaconda Command Prompt
* **Otter**: It is an autograder that we will be using for grading your notebooks.
* **grader.check('qn')**: This command runs test cases for nth question `qn` provided to you and displays the result. But these are not the only test cases, these are provided just for your basic testing. There will be some *hidden test cases* on which your answers will be tested as well.
* You are **not** allowed to edit any pre-defined variables and as per the instructions for every question you have to assign your answers to these variables.
* Write your answers in place of `...`


<h3>Submission Instructions</h3>


* There are two parts of this assignment. You should name both notebooks as **YourRollNumber_PartNumber.ipynb**. (e.g. this part should be named as *`2XXXXXXX_2.ipynb`*)
* You must submit a zip file containg just two notebooks and it must not has any folder or any other file. 
* The name of the zip file must be your `RollNumber_PA1`


Please run the cell below to import all the necessary libraries that you would require for this part

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

<h3>Context about the datasets</h3>

Given the rise in obesity we are aware that there is a huge chance that a person may be affected by the heart disease. In this assignment we are going to explore a dataset which contains information about individuals heart rate and oxygen saturation. In this part we are going to analyze the data and generate sufficent results from it.

Note: The output column is binary (0 or 1) and signifies whether the individual has had a heart attack (1 if they have had a heart attack).

**Question 0:** Please load the two dataset files inside the variables `heartDf` and `o2SaturationDf` respectively

```python
heartDf = pd.read_csv('heart.csv')
o2SaturationDf = pd.read_csv('o2Saturation.csv')
```

```python
grader.check("Q0")
```

Output:
```
Q0 results: All test cases passed!
```

**Question 1:** Find the shapes of both datasets and assign them to `heartDfShape` and `o2DfShape` respectively

```python
heartDfShape = heartDf.shape
o2DfShape = o2SaturationDf.shape
```

```python
grader.check("Q1")
```

Output:
```
Q1 results: All test cases passed!
```

**Question 2A:** Use a pandas function to display all characteristics (including min, max, quartiles, count, mean) of all numerical attributes of both dataframes. Store these in the variables `heartdf_description` & `o2df_description` respectively. 

```python
heartdf_description = heartDf.describe()
o2df_description = o2SaturationDf.describe()
```

```python
grader.check("Q2a")
```

Output:
```
Q2a results: All test cases passed!
```

**Question 2B:** Use a pandas function to display the number of occurances of all unique values of sex in the dataset. Store these in the variable `gender_distribution`. 

```python
gender_distribution = heartDf['sex'].value_counts() 
```

```python
grader.check("Q2B")
```

Output:
```
Q2B results: All test cases passed!
```

**Question 3A:** Prepare the two dataframes for merging by making the following column changes to both the dataframes. 
<h5>For Heart Dataset</h5>
<ul>
<li>exng -> exercise induced angina</li>
<li>caa -> number of major vessels</li>
<li>cp -> chest pain type</li>
<li>trtbps -> resting blood pressure</li>
<li>chol -> cholestrol levels</li>
<li>restecg -> resting electrocardiographic </li> 
<li>thalachh-> maximum heart rate</li>
<li>fbs-> fasting blood sugar > 120mg/dl</li>
<li>oldpeak -> old peak </li>
</ul>
<h5>For O2 Saturation Dataset </h5>
<ul>
<li>sat -> O2 saturation</li>
</ul>


```python
heartDf.rename(columns={'exng':'exercise induced angina','caa':'number of major vessels','cp':'chest pain type','trtbps':'resting blood pressure','chol':'cholesterol levels','restecg':'resting electrocardiographic','thalachh':'maximum heart rate','fbs':'fasting blood sugar > 120mg/dl','oldpeak':'old peak'}, inplace=True) 
o2SaturationDf.rename(columns={'sat':'O2 saturation'}, inplace=True)
```

```python
grader.check("Q3A")
```

Output:
```
Q3A results: All test cases passed!
```

**Question 3B:** Peform a merge operation so that we can better compare our results and reach an appropriate conclusion. 

```python
mergedDf = pd.concat([heartDf, o2SaturationDf], axis=1)
```

```python
grader.check("Q3B")
```

Output:
```
Q3B results: All test cases passed!
```

**Question 4:** Categorize the individuals based on their age group using the following groupings in the corresponding variable names with `loc`:
<ul>
<li>Below 40 years of age</li>
<li>In the range of 40-50 years of age</li>
<li>In the range of 50-60 years of age</li>
<li>In the range of 60-70 years of age</li>
<li>70 or above years of age</li>
</ul>

<b>Note: </b> Only the lower bound is inclusive in the above age group range. You are to store the results in the variables given in the cell below.

```python
belowForty = mergedDf.loc[mergedDf['age'] < 40]
ageBtw40_50 = mergedDf.loc[(mergedDf['age'] >= 40) & (mergedDf['age'] < 50)]
ageBtw50_60 = mergedDf.loc[(mergedDf['age'] >= 50) & (mergedDf['age'] < 60)]
ageBtw60_70 = mergedDf.loc[(mergedDf['age'] >= 60) & (mergedDf['age'] < 70)]
aboveSeventy = mergedDf.loc[mergedDf['age'] >= 70]
```

```python
grader.check("Q4")
```

Output:
```
Q4 results: All test cases passed!
```

**Question 5A:**  Group the data by 'sex' and find the average age for each group. <br/>

Store the result in the variable `mean_ages_by_sex`.

```python
mean_ages_by_sex = mean_ages_by_sex = mergedDf.groupby('sex')['age'].mean()
```

```python
grader.check("Q5A")
```

Output:
```
Q5A results: All test cases passed!
```

**Question 5B:**  Extract rows for patients whose maximum heart rate is within 10% of the mean heart rate across all patients. <br/>

Store the result in the variable `patients_with_normal_heartrate`.

```python
mean_hr = mergedDf['maximum heart rate'].mean()
patients_with_normal_heartrate = mergedDf.loc[(mergedDf['maximum heart rate'] >= mean_hr * 0.9) & (mergedDf['maximum heart rate'] <= mean_hr * 1.1)]
```

```python
grader.check("Q5B")
```

Output:
```
Q5B results: All test cases passed!
```

**Question 5C:**  For patients younger than 50, find the average cholesterol levels for those with a maximum heart rate above 160. <br/>

Store the result in the variable `mean_cholesterol_youngerPatients`.

```python
df_filtered = mergedDf.loc[(mergedDf['age'] < 50) & (mergedDf['maximum heart rate'] > 160)]
mean_cholesterol_youngerPatients = df_filtered['cholesterol levels'].mean()
```

```python
grader.check("Q5C")
```

Output:
```
Q5C results: All test cases passed!
```

**Question 5D:** Define a custom function that classifies a patient as `high risk` if their cholesterol level is above 240 and their `maximum heart rate` is below 120. Apply this function to each group of patients with different `chest pain type`s and count how many patients meet these high-risk conditions for each chest pain type. You will be using `.apply` for this task. <br/>

Store the result in the variable `high_risk_counts`. 

```python
def classify_high_risk(group):
    return group[(group['cholesterol levels'] > 240) & (group['maximum heart rate'] < 120)].shape[0]

```

```python
high_risk_counts = mergedDf.groupby('chest pain type').apply(classify_high_risk)
```

Output:
```
C:\Users\razah\AppData\Local\Temp\ipykernel_12852\2662984840.py:1: DeprecationWarning: DataFrameGroupBy.apply operated on the grouping columns. This behavior is deprecated, and in a future version of pandas the grouping columns will be excluded from the operation. Either pass `include_groups=False` to exclude the groupings or explicitly select the grouping columns after groupby to silence this warning.
  high_risk_counts = mergedDf.groupby('chest pain type').apply(classify_high_risk)
```

```python
grader.check("Q5D")
```

Output:
```
Q5D results: All test cases passed!
```

**Question 6:**  Is the average resting blood pressure of individuals in the (Above 70 age group) greater than the resting blood pressure of inviduals that are below 40 years old? <br/>
Store the result in `meanOfAboveSeventy` & `meanOfBelowForty` variables. <br/> Also the variable `ans` must contain a string such as a 'Yes' or 'No' as the test cases check for it. <br>
<b>Note: </b>You are not allowed to use the mean() method but rather calculate the mean mathematically.

```python
sum_above_seventy = aboveSeventy['resting blood pressure'].sum()
count_above_seventy = aboveSeventy['resting blood pressure'].count()
meanOfAboveSeventy = sum_above_seventy / count_above_seventy

sum_below_Forty = belowForty['resting blood pressure'].sum()
count_below_Forty = belowForty['resting blood pressure'].count()
meanOfBelowForty = sum_below_Forty / count_below_Forty
q6ans = ans = 'Yes' if meanOfAboveSeventy > meanOfBelowForty else 'No'
```

```python
grader.check("Q6")
```

Output:
```
Q6 results: All test cases passed!
```

**Question 7:** What is the most occuring value for oxygen saturation in individuals that fall into the age group of 50-60 ? <br>
Store the value in the variable `frequentO2Saturation`.

```python
frequentO2Saturation = ageBtw50_60['O2 saturation'].mode().values[0]
```

```python
grader.check("Q7")
```

Output:
```
Q7 results: All test cases passed!
```

**Question 8:** Find the most prevalent type of chest pain for all age groups mentioned above and store it in the corresponding variable. Also, find the most prevalent type of chest pain across all the individuals and store your result in `mostObservedChestPain`.

```python
belowFortyChestPain = belowForty['chest pain type'].mode().values[0]
ageBtw40_50ChestPain = ageBtw40_50['chest pain type'].mode().values[0]
ageBtw50_60ChestPain = ageBtw50_60['chest pain type'].mode().values[0]
ageBtw60_70ChestPain = ageBtw60_70['chest pain type'].mode().values[0]
aboveSeventyChestPain = aboveSeventy['chest pain type'].mode().values[0]
mostObservedChestPain = mergedDf['chest pain type'].mode().values[0]
```

```python
grader.check("Q8")
```

Output:
```
Q8 results: All test cases passed!
```

**Question 9:** Calculate the <b>median</b> cholestrol levels for all the age groups.<br>
<b>Note: </b> Here the variable `listOfDf` contains all the age groups that you have previously calculated in question <b>6</b>.

```python
listOfDf = listOfDf = [belowForty, ageBtw40_50, ageBtw50_60, ageBtw60_70, aboveSeventy]
medianOfAllCholLevels = [df['cholesterol levels'].median() for df in listOfDf]
```

```python
grader.check("Q9")
```

Output:
```
Q9 results: All test cases passed!
```

**Question 10:**  Using a pivot table, show the gender distribution across dfferent chest pain types. Store the result in `gender_count`.

```python
gender_count = mergedDf.pivot_table(index='chest pain type', columns='sex', aggfunc='size', fill_value=0)
```

```python
grader.check("Q10")
```

Output:
```
Q10 results: All test cases passed!
```

**Question 11:**  Calculate the percentage of males (1) and females (0) who are diagnosed with exercise induced angina. Round the result to three decimal places and store the result in `male_prop` and `female_prop` respectively<br/>
<b>Note:</b> 1 means they have exercise induced angina whereas 0 means they do not. 

```python
total_males = mergedDf[mergedDf['sex'] == 1].shape[0]
total_females = mergedDf[mergedDf['sex'] == 0].shape[0]

males_with_angina = mergedDf[(mergedDf['sex'] == 1) & (mergedDf['exercise induced angina'] == 1)].shape[0]
females_with_angina = mergedDf[(mergedDf['sex'] == 0) & (mergedDf['exercise induced angina'] == 1)].shape[0]

male_prop = round((males_with_angina / total_males) * 100, 3)
female_prop = round((females_with_angina / total_females) * 100, 3)
```

```python
grader.check("Q11")
```

Output:
```
Q11 results: All test cases passed!
```

**Question 12:**  What is the maximum heart rate across the different age groups ? 
Store the max heart rates obtained in the variable `heartRates`.

```python
heartRates = [belowForty['maximum heart rate'].max(), ageBtw40_50['maximum heart rate'].max(), ageBtw50_60['maximum heart rate'].max(), ageBtw60_70['maximum heart rate'].max(), aboveSeventy['maximum heart rate'].max()]
```

```python
grader.check("Q12")
```

Output:
```
Q12 results: All test cases passed!
```

**Question 13A:**  How many people have cholestrol values greater than the threshold value given below ? Store it in the variable `numberOfPeople`. 
<br>
<ul>
<li>Cholestrol Threshold Level: 300</li>
</ul>

Also calculate the probability of these people getting a heart attack and store the result, rounded off to three decimal places, in `probability`.


```python
cholesterol_threshold = 300

people_with_high_cholesterol = mergedDf[mergedDf['cholesterol levels'] > cholesterol_threshold]
numberOfPeople = people_with_high_cholesterol.shape[0]

if numberOfPeople > 0:
    probability = round((people_with_high_cholesterol['output'].sum() / numberOfPeople), 3)
else:
    probability = 0.0
```

```python
grader.check("Q13A")
```

Output:
```
Q13A results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**Question 13B:**  Do you think people with such cholesetrol levels have a higher probability of getting an heart attack than those with lower cholesterol levels? Justify this claim by providing a valid reason.

```python
low_cholesterol_people = mergedDf[mergedDf['cholesterol levels'] <= cholesterol_threshold]

# Count the number of people with cholesterol levels <= 300
numberOfLowCholesterolPeople = low_cholesterol_people.shape[0]

# Calculate the probability of a heart attack for people with cholesterol <= 300
low_chol_probability = round(low_cholesterol_people['output'].sum() / numberOfLowCholesterolPeople, 3) if numberOfLowCholesterolPeople > 0 else 0

# Display the result
print("Probability of heart attack for people with cholesterol <= 300:", low_chol_probability)
```

Output:
```
Probability of heart attack for people with cholesterol <= 300: 0.536
```

As seen above, the probability of getting a heartattack despite having a lower CHolestrol Level is nearly the same. This data makes the claim that having a cholestrol above 300 does not make you more sucpetible to having a heart attack. The reasons for this outcome are many.

Firstly, Cholestrol alone does not play a massive factor in determening if a person is susceptible to a heart attack. Genetics, mental health, Blood Pressure and Lifestyle also play a big role in that.

Secondly, the data we have fore people with high cholestrol is very low. Only 43 people with high cholestrol and 265 with low cholesstrol, so the data is unevenly balanced.

Lastly, it's entirely possible in the small dataset, for outliers or exceptions to exist which would introduce a large error margin due to the small size of the data.

However, going off of the data collected, assuming it's correct and accurate data, we come to the conclusion that having high cholestrol doesn't increase the probability of a heartattack. (See reason 1)

<!-- END QUESTION -->

**This marks the end of Assignment 1 Part 2 :) <br> Hooray! You made it to the end now give your self a pat on the back!**

Tests for Q0 (points: 1)
```
>>> assert heartDf.shape == (309, 14)
>>> assert o2SaturationDf.shape == (309, 1)
```


Tests for Q1 (points: 1)
```
>>> assert heartDfShape == (309, 14)
>>> assert o2DfShape == (309, 1)
```


Tests for Q10 (points: 3)
```
>>> assert gender_count.shape == (4, 2)
>>> assert gender_count[0].sum() == 101
>>> assert gender_count[1].sum() == 208
```


Tests for Q11 (points: 4)
```
>>> assert male_prop == 37.019
>>> assert female_prop == 21.782
```


Tests for Q12 (points: 3)
```
>>> assert heartRates == [202.0, 194.0, 195.0, 179.0, 162.0]
```


Tests for Q13A (points: 3)
```
>>> assert numberOfPeople == 43
>>> assert probability == 0.535
```


Tests for Q2B (points: 2)
```
>>> assert gender_distribution[1] == 208
>>> assert gender_distribution[0] == 101
```


Tests for Q2a (points: 2)
```
>>> assert heartdf_description.shape == (8, 14)
>>> assert o2df_description.shape == (8, 1)
```


Tests for Q3A (points: 3)
```
>>> assert 'exercise induced angina' in heartDf.columns
>>> assert 'number of major vessels' in heartDf.columns
>>> assert 'chest pain type' in heartDf.columns
>>> assert 'fasting blood sugar > 120mg/dl' in heartDf.columns
>>> assert 'resting blood pressure' in heartDf.columns
>>> assert 'old peak' in heartDf.columns
>>> assert 'cholesterol levels' in heartDf.columns
>>> assert 'resting electrocardiographic' in heartDf.columns
>>> assert 'maximum heart rate' in heartDf.columns
>>> assert 'O2 saturation' in o2SaturationDf.columns
```


Tests for Q3B (points: 2)
```
>>> assert mergedDf.shape == (309, 15)
>>> assert 'O2 saturation' in mergedDf.columns
```


Tests for Q4 (points: 5)
```
>>> assert belowForty.shape == (16, 15)
>>> assert ageBtw40_50.shape == (72, 15)
>>> assert ageBtw50_60.shape == (131, 15)
>>> assert ageBtw60_70.shape == (80, 15)
>>> assert aboveSeventy.shape == (10, 15)
```


Tests for Q5A (points: 2)
```
>>> assert mean_ages_by_sex[0] == 55.79207920792079
>>> assert mean_ages_by_sex[1] == 53.77403846153846
```


Tests for Q5B (points: 3)
```
>>> assert patients_with_normal_heartrate.shape == (143, 15)
```


Tests for Q5C (points: 2)
```
>>> assert mean_cholesterol_youngerPatients == 232.3653846153846
```


Tests for Q5D (points: 3)
```
>>> assert high_risk_counts[0] == 15
>>> assert high_risk_counts[1] == 1
>>> assert high_risk_counts[2] == 2
>>> assert high_risk_counts[3] == 0
```


Tests for Q6 (points: 4)
```
>>> assert q6ans.strip().lower() == 'yes'
>>> assert meanOfAboveSeventy == 135.8
>>> assert meanOfBelowForty == 125.5
```


Tests for Q7 (points: 2)
```
>>> assert frequentO2Saturation == 97.5
```


Tests for Q8 (points: 5)
```
>>> assert belowFortyChestPain == 2
>>> assert ageBtw40_50ChestPain == 0
>>> assert ageBtw50_60ChestPain == 0
>>> assert ageBtw60_70ChestPain == 0
>>> assert aboveSeventyChestPain == 0
>>> assert mostObservedChestPain == 0
```


Tests for Q9 (points: 4)
```
>>> assert medianOfAllCholLevels == [207.0, 235.5, 239.0, 252.5, 267.0]
```
