```python
# Initialize Otter
import otter
grader = otter.Notebook("26100319_1.ipynb")
```

# **Assignment 1 (Part 1)**

<h3>Assignment Instructions</h3>


* **How to install 'otter'**: Run `pip install otter-grader` in your Anaconda Command Prompt
* **Otter**: It is an autograder that we will be using for grading your notebooks.
* **Initialize Otter**: Run the cell above to initialize otter. Replace `2XXXXXX_1.ipynb` with your updated filename inside the grader variable above. 
* **grader.check('qn')**: This command runs test cases for nth question `qn` provided to you and displays the result. But these are not the only test cases, these are provided just for your basic testing. There will be some *hidden test cases* on which your answers will be tested as well.
* You are **not** allowed to edit any pre-defined variables and as per the instructions for every question you have to assign your answers to these variables.
* Write your answers in place of `...`


<h3>Submission Instructions</h3>


* There are two parts of this assignment. You should name both notebooks as **YourRollNumber_PartNumber.ipynb**. (e.g. this part should be named as *`2XXXXXXX_1.ipynb`*)
* You must submit a zip file containg just two notebooks and it must not has any folder or any other file. 
* The name of the zip file must be your `RollNumber_PA1`


Please run the cell below to import all the necessary libraries that you would require for this part

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

<h3>Context about the dataset</h3>


The titanic dataset is a record of passengers aboard the Titanic and their fates (whether or not they survived the incident). This includes socio-economic class divide between passengers, their ages, names, fares and even their passenger IDs. The principal source for data about Titanic passengers is the Encyclopedia Titanica. The datasets used here were begun by a variety of researchers. One of the original sources is Eaton & Haas (1994) Titanic: Triumph and Tragedy, Patrick Stephens Ltd, which includes a passenger list created by many researchers and edited by Michael A. Findlay.

As we explore this data from titanic, it is your job to answer some pressing questions regarding the various age groups aboard, their economic statuses and survival rates in each gender. Following are some details about the variable to help you understand the data better. <br>

<h4>Variable Descriptions</h4>

* Pclass ->   Passenger Class  (1 = 1st; 2 = 2nd; 3 = 3rd) <br>
* Survived ->  Did the passenger survive? (0 = No; 1 = Yes) <br>
* Name ->  name <br>
* Sex ->  sex  <br>
* Age ->  age  <br>
* SibSp ->  Number of Siblings/Spouses Aboard <br>
* Parch ->  Number of Parents/Children Aboard <br>
* Ticket ->   Ticket Number <br>
* Fare ->   Passenger Fare (British pound) <br>
* Cabin ->   Cabin <br>
* Embarked ->   Port of Embarkation (C = Cherbourg; Q = Queenstown; S = Southampton) <br>

**Question 0:** Load the titanic dataset file into the variable `titanic_df`

```python
titanic_df = pd.read_csv('titanic.csv')
```

```python
grader.check("Q0")
```

Output:
```
Q0 results: All test cases passed!
```

**Question 1:** Using the describe function, find the mean fare  and store it into `mean_fare`, rounded off to three decimal places. You cannot directly call the .mean() function.

```python
fare_stats = titanic_df['Fare'].describe()
mean_fare = round(fare_stats['mean'], 3)
```

```python
grader.check("Q1")
```

Output:
```
Q1 results: All test cases passed!
```

**Question 2:** Drop the `Cabin` and `Family_Size` column inside the dataframe. `titanic_df` should no longer contain these columns.

```python
titanic_df = titanic_df.drop(columns=['Cabin', 'Family_Size'])
```

```python
grader.check("Q2")
```

Output:
```
Q2 results: All test cases passed!
```

**Question 3:** What percentage of people aboard the ship were females? Store the result rounded off to three decimal places in `female_percentage`.

```python
num_females = titanic_df[titanic_df['Sex'] == 'female'].shape[0]

total_passengers = titanic_df.shape[0]

female_percentage = round((num_females / total_passengers) * 100, 3)
```

```python
grader.check("Q3")
```

Output:
```
Q3 results: All test cases passed!
```

**Question 4:** Set the `PassengerID` of passengers as the index for this dataset.

```python
titanic_df = titanic_df.set_index('PassengerId')
```

```python
grader.check("Q4")
```

Output:
```
Q4 results: All test cases passed!
```

**Question 5:** Convert all column names to lowercase and remove any leading and trailing whitespaces. Simultaneously, replace any non-leading and non-trailing spaces within the column names with an underscore. You are not allowed to hardcode the column names.

```python
titanic_df.columns = titanic_df.columns.str.lower().str.strip().str.replace(' ', '_')
```

```python
grader.check("Q5")
```

Output:
```
Q5 results: All test cases passed!
```

**Question 6:** Find the female passenger that had the highest fare and store the result in `highest_fare`. Furthermore, find the male passenger that had the lowest non-zero fare and store it in `lowest_fare`. <br>
Note: You are not meant to store the value only but the entire row entry.

```python
highest_fare = titanic_df[titanic_df['sex'] == 'female'].nlargest(1, 'fare')
lowest_fare = titanic_df[(titanic_df['sex'] == 'male') & (titanic_df['fare'] > 0)].nsmallest(1, 'fare')
```

```python
grader.check("Q6")
```

Output:
```
Q6 results: All test cases passed!
```

**Question 7:** Write a function `CalculateSurvival(data)` that takes in a dataframe and returns the survival percentage rounded off to three decimal places.

```python
def CalculateSurvivalRate(data):
    total_passengers = data.shape[0]
    survivors = data[data['survived'] == 1].shape[0]
    survival_percentage = (survivors / total_passengers) * 100
    
    return round(survival_percentage, 3)
```

```python
grader.check("Q7")
```

Output:
```
Q7 results: All test cases passed!
```

**Question 8A:** Add a column called `family` to the dataframe that shows true for a passenger whose family members (siblings, spouses, parents, or children) were on board and false otherwise.

```python
titanic_df["family"] = (titanic_df['sibsp'] + titanic_df['parch']) > 0
```

```python
grader.check("Q8A")
```

Output:
```
Q8A results: All test cases passed!
```

**Question 8B:** Using the `CalculateSurvivalRate` function, find the survival percentage of the passengers who had family on board (siblings, spouses, children or parents) and store the result in `survival_rate_with_family`.

```python
survival_rate_with_family = CalculateSurvivalRate(titanic_df[titanic_df['family'] == True])
```

```python
grader.check("Q8B")
```

Output:
```
Q8B results: All test cases passed!
```

**Question 8C:** Using the `CalculateSurvivalRate` function, find the survival percentage of the passengers who did not have any family on board (no siblings, spouses, children and parents) and store the result in `survival_rate_without_family`.

```python
survival_rate_without_family = CalculateSurvivalRate(titanic_df[titanic_df['family'] == False])
```

```python
grader.check("Q8C")
```

Output:
```
Q8C results: All test cases passed!
```

**Question 8D:** Based on your findings above, were the passengers who had family more likely to survive than those who did not? Answer as 'Yes' or 'No' and store it in `ans`.

```python
ans = 'Yes'
```

```python
grader.check("Q8D")
```

Output:
```
Q8D results: All test cases passed!
```

**Question 9:** Use `groupby` to find the survival percentage of males and females and store them in `male_survival_percentage` and `female_survival_percentage` respectively. Round off your answers to three decimal places. Not using `groupby` will result in a zero for this question.

```python
survival_by_gender = titanic_df.groupby('sex')['survived'].mean() * 100

male_survival_percentage = round(survival_by_gender['male'], 3)
female_surival_percentage = round(survival_by_gender['female'], 3)

```

```python
grader.check("Q9")
```

Output:
```
Q9 results: All test cases passed!
```

**Question 10:** Using the pandas iloc function, extract the rows with passengers aged between 5 to 20 and fare above 20. Store this result in `ilocresult`. 

```python
filtered_df = titanic_df[(titanic_df['age'] >= 5) & (titanic_df['age'] <= 20) & (titanic_df['fare'] > 20)]

ilocresult = filtered_df.iloc[:, [titanic_df.columns.get_loc('age'), titanic_df.columns.get_loc('fare')]]
```

```python
grader.check("Q10")
```

Output:
```
Q10 results: All test cases passed!
```

**Question 11A:** Replace the integer values in the `Pclass` columns with corresponding string data types. For example, replace 1 with 'Upper', 2 with 'Middle' and 3 with 'Lower'.

```python
titanic_df['pclass'] = titanic_df['pclass'].replace({1: 'Upper', 2: 'Middle', 3: 'Lower'})
```

```python
grader.check("Q11A")
```

Output:
```
Q11A results: All test cases passed!
```

**Question 11B:** Using a pivot table, represent the male and female passengers that survived in each of the three passenger classes. Store the result in `class_count`.

```python
class_count = titanic_df.pivot_table(index='pclass', columns='sex', values='survived', aggfunc='sum', fill_value=0)
```

```python
grader.check("Q11B")
```

Output:
```
Q11B results: All test cases passed!
```

**Well done on completing the first part of your assignment. Take a break and move onto the next!**

Tests for Q0 (points: 1)
```
>>> assert titanic_df.shape == (891, 14)
```


Tests for Q1 (points: 3)
```
>>> assert mean_fare == 32.204
```


Tests for Q10 (points: 3)
```
>>> assert ilocresult.shape == (53, 2)
>>> assert ilocresult['age'].median() == 15.0
>>> assert ilocresult['fare'].median() == 34.375
```


Tests for Q11A (points: 2)
```
>>> assert titanic_df['pclass'].value_counts()['Upper'] == 216
>>> assert titanic_df['pclass'].value_counts()['Lower'] == 491
>>> assert titanic_df['pclass'].value_counts()['Middle'] == 184
```


Tests for Q11B (points: 4)
```
>>> assert list(class_count['female']) == [72, 70, 91]
>>> assert list(class_count['male']) == [47, 17, 45]
```


Tests for Q2 (points: 3)
```
>>> assert titanic_df.shape == (891, 12)
>>> assert 'Cabin' not in titanic_df.columns
>>> assert 'Family_Size' not in titanic_df.columns
```


Tests for Q3 (points: 3)
```
>>> assert female_percentage == 35.241
```


Tests for Q4 (points: 2)
```
>>> assert 'PassengerId' not in titanic_df.columns
>>> assert titanic_df.index.name == 'PassengerId'
>>> assert titanic_df.shape == (891, 11)
```


Tests for Q5 (points: 3)
```
>>> assert list(titanic_df.columns) == ['age', 'embarked', 'fare', 'name', 'parch', 'pclass', 'sex', 'sibsp', 'survived', 'ticket', 'title']
```


Tests for Q6 (points: None)
```
>>> assert highest_fare.index[0] == 259
>>> assert highest_fare['fare'].values[0] == 512.3292
>>> assert lowest_fare['fare'].values[0] == 4.0125
>>> assert lowest_fare.index[0] == 379
```


Tests for Q7 (points: 4)
```
>>> assert CalculateSurvivalRate(titanic_df.loc[(titanic_df['sibsp'] > 0) | (titanic_df['parch'] > 0)]) == 50.565
>>> assert CalculateSurvivalRate(titanic_df.loc[(titanic_df['sibsp'] == 0) & (titanic_df['parch'] == 0)]) == 30.354
```


Tests for Q8A (points: 2)
```
>>> assert titanic_df['family'].value_counts()[True] == 354
>>> assert titanic_df['family'].value_counts()[False] == 537
```


Tests for Q8B (points: 2)
```
>>> assert survival_rate_with_family == 50.565
```


Tests for Q8C (points: 2)
```
>>> assert survival_rate_without_family == 30.354
```


Tests for Q8D (points: 2)
```
>>> assert ans.strip().lower() == 'yes'
```


Tests for Q9 (points: 3)
```
>>> assert male_survival_percentage == 18.891
>>> assert female_surival_percentage == 74.204
```
